# LIN 373 Machine Learning Toolbox for Text Analysis, Spring 2019

# Homework 3 - due Thursday Mar 12 2020 at 2:00pm

For this homework you will hand in (upload) to canvas:
- JUST this notebook renamed ``hw3_YourEID.ipynb``

__Before submitting__, please reset your kernel and rerun everything from the beginning (`Kernel` >> `Restart and Run All`) an ensure your code outputs the correct answer. I

A perfect solution for this homework is worth **100** points. For programming tasks, make sure that your code can run using Python 3.5+. If you cannot complete a problem, include as much pseudocode as possible for partial credit. However, make sure it does not have any output errors. **If there are any output errors, half of the points for that problem will be automatically deducted.**

Collaboration: you are free to discuss the homework assignments with other students and work towards solutions together.  However, all of the code you write must be your own! There is a thread on Piazza where you can look for a study group: https://piazza.com/class/k5o768x0bq6ai?cid=10

Review extension and academic dishonesty policy here: https://jjessyli.github.io/courses/lin373#extension-policy

For typing up your answers to non-programming problems, information can be found about Markdown cells for Jupyter Notebooks here: https://jupyter-notebook.readthedocs.io/en/stable/examples/Notebook/Working%20With%20Markdown%20Cells.html



### Please list any collaborators here:
- 


## Problem 1: Logistic Regression

* **(a)** (5 points) In logistic regression, we know that when there are two classes an example can belong to ($y\in\{0,1\}$), for each example ${\bf x}$ with M features,
    $$
    p(y=1|x_1, x_2, ..., x_M) = \frac{1}{1+\exp{ (-\sum_{j=1}^M w_jx_j-b) }}
    $$
    $$
    p(y=0|x_1, x_2, ..., x_M) = 1-p(y=1|x_1, x_2, ..., x_M)
    $$
Generally, how many parameters does logistic regression estimate? Please explain any variables in your answer. 


<!-- grade_id: answer-1a, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

The logistic regression usually has many different parameters. It is dependent on how features there are. For example, usually for every feature, you have an associated weight, and then you also have a bias term. 

* **(b)** (10 points) Given a dataset of $N$ examples, briefly describe how these parameters are estimated. You do not need to lay out specific mathematical derivations, rather, provide the name(s) of the process(es) and give a few sentences describing the process. You can assume that all these examples will be used for training.

<!-- grade_id: answer-1b, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE
Usually, you will start off with certain weights for each feature. Say we have four different features, and a bias term. We could set them all to zero. Then, we can multiple the amount of features of a specific data set by these weigths and then put it through a sigmoid function to determine the probability of it being in a certain class. For example, say after multiplying the vectors,adding the bias, and putting it through a sigmoid function we get a number like 0.51. This means that are model has predicted that the data set is of class 1. Say our model is wrong though, and it is actually in class 0. Then we can use the cross-entropy lost function to determine how far we are off. What we want to do is minimize the value we get for the cross-entropy function. To minimize it we will use gradient descent to get the optimal weigths for each feature that minimize the cross entropy function. This will take many iterations as each interation nudges the weights closer to minimizing the cors entropy function


* **(c)** (10 points) Logistic regression can be extended to classify $K$ classes instead of only 2. When $y$ can be one of K classes, the conditional probability of $y$ being of a particular class $k$ is:
 $$
 p(y=k|x_1,x_2,...,x_M) = \frac{\exp(\sum_{j=1}^M w_{jk}x_j+b_k)}{\sum_{k'=1}^K\exp(\sum_{j=1}^M w_{jk'}x_j+b_{k'})}
 $$
How many parameters are we estimating in this case?  Please explain any variables in your answer. 

<!-- grade_id: answer-1c, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE
The multinomial logistic regression is similar to the binary logistic regression. Some exceptions are that we need separate weight vectors for each of the k classes we have, and that we use a softmax function instead of the sigmoid which effect the loss function and gradient function.

## Problem 2: Regularization and evaluation

Recall that with Ridge Regression, we add a _regularization_ term to the log likelihood:
$$
    l({\bf w}) = \log \prod_{i=1}^N{ p(y^{(i)}|{\bf x^{(i)},w})}-\lambda||{\bf w}||_2^2
$$

* **(a)** (10 points) What is the goal of this regularization term? Describe how it works in one sentence. 

<!-- grade_id: answer-2a, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE
The goal is to penalize weights that are really big which could cause problems. This is done by adding a bias term to the equation which makes it less sensitive to changes in it parameters. T 

* **(b)** (15 points) Your friend David uses cross validation to pick $\lambda$ for his application of logistic regression as follows. He sets aside 20\% of his data as a test set, and then runs cross validation on his training data to get cross validation accuracy with multiple values of $\lambda$. He also trains a model on the entire training data for each value of $\lambda$ and computes test accuracy on his held out data. He gets the following results:

| $\lambda$ | Cross-validation accuracy | Test accuracy |
|:---------:|:-------------------------:|:-------------:|
|     0     |            0.7            |      0.65     |
|     1     |            0.75           |      0.7      |
|     10    |            0.8            |      0.63     |
|    100    |            0.7            |      0.6      |

* **(b-i)** What is the optimal $\lambda$ in terms of test accuracy? Use the variable `optimal_test_accuracy`

<!-- grade_id: cell-20d754030dda7fb6, points: 0, grade: false, solution: true -->
```python
# initialize the variable optimal_test_accuracy with the correct answer!
# YOUR CODE HERE
optimal_test_accuracy = 1
#raise NotImplementedError()
```

<!-- grade_id: cell-a549a195010745fe, points: 5, grade: true, solution: false -->
```python
## Do not edit this cell - hidden autograder test
```

* **(b-ii)** Is this the same $\lambda$ that would be chosen by cross validation? Use the variable `optimal_cross_validation`

<!-- grade_id: cell-53232bb51fea7452, points: 0, grade: false, solution: true -->
```python
# initialize the variable optimal_cross_validation with the correct answer!

optimal_cross_validation = 10
#raise NotImplementedError()


```

<!-- grade_id: cell-3f810ebb88eb1947, points: 5, grade: true, solution: false -->
```python
## Do not edit this cell - hidden autograder test
```

* **(b-iii)** David wants to inform the public about his model and his results. What should he report as the test accuracy of his method (logistic regression with regularization parameter $\lambda$) on this dataset? Use the variable `reported_accuracy`. 

<!-- grade_id: cell-f31fd5f701eca413, points: 0, grade: false, solution: true -->
```python
# initialize the variable reported_accuracy with the correct answer!
reported_accuracy = 0.63
#raise NotImplementedError()

```

<!-- grade_id: cell-17e26c3376faee4b, points: 5, grade: true, solution: false -->
```python
## Do not edit this cell - hidden autograder test
```

## Problem 3: Logistic Regression in sklearn

Logistic regression, it turns out, is a natural "cousin" to Naive Bayes; specifically it is the "discriminative" variant. First, let's review the code for the sklearn version of [Naive Bayes (sentiment analysis notebook)](https://utexas.instructure.com/courses/1267031/files/folder/Code-Data/10-oop-nb). Note that we define a random state for , so that the answers will be consistent across everyone (it should be 0.8175)! 

```python
%%time
from sklearn.model_selection import train_test_split
from sklearn.naive_bayes import MultinomialNB
import numpy as np
from sklearn.datasets import load_files
from sklearn import metrics
from sklearn.feature_extraction.text import CountVectorizer
from tqdm import tqdm


# load data using sklearn.datasets.load_files
dataset = load_files("movie-reviews/")


# split the data into train and test
docs_train, docs_test, y_train, y_test = train_test_split(
dataset.data, dataset.target, test_size = 0.2, random_state = 3)

# vectorize the training data
vectorizer = CountVectorizer(stop_words = "english")
X_train = vectorizer.fit_transform(docs_train)

# fit the model with the training data
nb_model = MultinomialNB()
nb_model.fit(X_train, y_train)

# vectorize the test data and predict 
X_test = vectorizer.transform(docs_test)
y_hat_nb = nb_model.predict(X_test)

# get accuracy score
accuracy_score_nb = metrics.accuracy_score(y_test, y_hat_nb)

# print out some data
print("nb_model accuracy is : ", accuracy_score_nb)
print("classification report:\n", metrics.classification_report(y_test, y_hat_nb))

```

Output:
```
nb_model accuracy is :  0.8175
classification report:
               precision    recall  f1-score   support

           0       0.81      0.82      0.82       197
           1       0.82      0.82      0.82       203

    accuracy                           0.82       400
   macro avg       0.82      0.82      0.82       400
weighted avg       0.82      0.82      0.82       400

Wall time: 30.1 s
```

```python
%%time
assert(accuracy_score_nb == 0.8175)
## If this throws and error, pleast contact the TA immediately!
```

Output:
```
Wall time: 0 ns
```

**(a)** (10 points) Using the code for NaiveBayes, implement [logistic regression model provided by sklearn](https://scikit-learn.org/stable/modules/generated/sklearn.linear_model.LogisticRegression.html). How do the two models compare in terms of predictive performance?

<!-- grade_id: cell-6fdc13b1e29e1fad, points: 0, grade: false, solution: true -->
```python
%%time
from sklearn.model_selection import train_test_split
from sklearn.linear_model import LogisticRegression
import numpy as np
from sklearn.datasets import load_files
from sklearn import metrics
from sklearn.feature_extraction.text import CountVectorizer
from tqdm import tqdm


# load data using sklearn.datasets.load_files
dataset = load_files("movie-reviews/")

# split the data into train and test
docs_train, docs_test, y_train, y_test = train_test_split(
dataset.data, dataset.target, test_size = 0.2, random_state = 3)

# vectorize the training data
vectorizer = CountVectorizer(stop_words = "english")
X_train = vectorizer.fit_transform(docs_train)


# fit the model with the training data
logmodel = LogisticRegression() #solver = 'lbfgs'
logmodel.fit(X_train, y_train)

# vectorize the test data and predict 
X_test = vectorizer.transform(docs_test)
y_hat_lr = logmodel.predict(X_test)

# get accuracy score
accuracy_score_lr = metrics.accuracy_score(y_test, y_hat_lr)

#raise NotImplementedError()

print("lr_model accuracy is : ", metrics.accuracy_score(y_test, y_hat_lr))
print("classification report:\n", metrics.classification_report(y_test, y_hat_lr))

```

Output:
```
C:\Users\Natha\Anaconda3\lib\site-packages\sklearn\linear_model\logistic.py:432: FutureWarning: Default solver will be changed to 'lbfgs' in 0.22. Specify a solver to silence this warning.
  FutureWarning)
lr_model accuracy is :  0.8475
classification report:
               precision    recall  f1-score   support

           0       0.87      0.81      0.84       197
           1       0.83      0.88      0.85       203

    accuracy                           0.85       400
   macro avg       0.85      0.85      0.85       400
weighted avg       0.85      0.85      0.85       400

Wall time: 3.3 s
```

```python
print(type(y_test)) 
```

Output:
```
<class 'numpy.ndarray'>
```

<!-- grade_id: cell-c18b15a3268013bb, points: 4, grade: true, solution: false -->
```python
## Do not edit this cell - autograder test
assert(accuracy_score_lr == 0.8475)
print("visable test cases passed!")
```

Output:
```
visable test cases passed!
```

<!-- grade_id: cell-1225245cd0051265, points: 0, grade: false, solution: true -->
```python
# Which model is better?
# initialize the variable better_basic_model with either "nb" or "lr"
better_basic_model="lr"
# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-5b0571001518b504, points: 1, grade: true, solution: false -->
```python
## Do not edit this cell - hidden autograder test
```

**(b)** (10 points) Instead of performing just one train/test split, try performing cross validation. [See here for documentation on doing this](https://scikit-learn.org/stable/modules/cross_validation.html). Now again compare NB and LR using 10-fold cross-validation. Use random_state = 3

<!-- grade_id: cell-b23dc25a2ffb21eb, points: 0, grade: false, solution: true -->
```python
%%time
from sklearn.model_selection import KFold
from sklearn.model_selection import cross_val_score
# use random_state = 3
# YOUR CODE HERE



vectorizer = CountVectorizer(stop_words = "english")
XX_train = vectorizer.fit_transform(dataset.data)
yy_train = dataset.target

logmodel = LogisticRegression(random_state=3)
nb_xval_scores = cross_val_score(nb_model, X_train, y_train, cv=10)
lr_xval_scores = cross_val_score(logmodel, X_train, y_train, cv=10)

kf = KFold(n_splits=10, random_state=3)

x = np.array(dataset.data)
y = np.array(dataset.target)

# models
logmodel = LogisticRegression()

### xx = np.array([1,2,3,4,5])
### print(xx[[2,3]]) ##3pretty cool feature for numpy



'''
for train, test in kf.split(dataset.data):
    print(test)
    docs_train, docs_test, y_train,y_test = x[train], x[test],y[train],y[test]
    vectorizer = CountVectorizer(stop_words = "english")
    X_train = vectorizer.fit_transform(docs_train)
    logmodel.fit(X_train, y_train)
    X_test = vectorizer.transform(docs_test)
    y_hat_lr = logmodel.predict(X_test)
    accuracy_score_lr = metrics.accuracy_score(y_test, y_hat_lr)
    nb_xval_scores.append(accuracy_score_lr)
    lr_xval_score = np.append(lr_xval_scores, [accuracy_score_lr])
    
print(lr_xval_score.mean())   
'''

print("Accuracy for Logistic Regression: %0.3f (+/- %0.2f)" % (lr_xval_scores.mean(), lr_xval_scores.std() * 2))
print("Accuracy for Naive Bayes: %0.2f (+/- %0.3f)" % (nb_xval_scores.mean(), nb_xval_scores.std() * 2))

```

Output:
```
C:\Users\Natha\Anaconda3\lib\site-packages\sklearn\linear_model\logistic.py:432: FutureWarning: Default solver will be changed to 'lbfgs' in 0.22. Specify a solver to silence this warning.
  FutureWarning)
Accuracy for Logistic Regression: 0.823 (+/- 0.06)
Accuracy for Naive Bayes: 0.80 (+/- 0.045)
Wall time: 5.62 s
```

<!-- grade_id: cell-bfcd9ca91b0ae335, points: 9, grade: true, solution: false -->
```python
## Do not edit this cell - autograder test
assert(round(lr_xval_scores.mean(),3) == 0.847)
assert(round(nb_xval_scores.mean(), 3) == 0.81)
print("visable test cases passed!")
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-15-9b075935edbd> in <module>
      1 ## Do not edit this cell - autograder test
----> 2 assert(round(lr_xval_scores.mean(),3) == 0.847)
      3 assert(round(nb_xval_scores.mean(), 3) == 0.81)
      4 print("visable test cases passed!")

NameError: name 'lr_xval_scores' is not defined
```

<!-- grade_id: cell-17788de9ee0fdc4d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
print(nb_xval_scores)
my_array= np.array(nb_xval_scores)
print(my_array.mean())
print(sum(nb_xval_scores)/10)

```

<!-- grade_id: cell-7facf3a5cd38d763, points: 1, grade: true, solution: false -->
```python
## Do not edit this cell - hidden autograder test
```

**(c)** (5 points) Now, let's modify the featurres in our classifier! The first thing we're trying is to have our classifier not only include unigrams, but also bigrams. Modify your `CountVectorizer` to include bigrams and use logistic regression as your classifier. Did the results on 10-fold cross validation improve? (Make sure to use the same folds as before, so the results are comparable!)

<!-- grade_id: cell-81286f74a36745a6, points: 0, grade: false, solution: true -->
```python
%%time
# include unigrams and bigrams


vectorizer = CountVectorizer(stop_words = "english", ngram_range = (1,2))
XX_train = vectorizer.fit_transform(dataset.data)
yy_train = dataset.target

logmodel = LogisticRegression(random_state=3)
nb_binary_scores = cross_val_score(nb_model, XX_train, yy_train, cv=10)
lr_binary_scores = cross_val_score(logmodel, XX_train, yy_train, cv=10)





print("Accuracy for Logistic Regression: %0.3f (+/- %0.2f)" % (lr_binary_scores.mean(), lr_binary_scores.std() * 2))
print("Accuracy for Naive Bayes: %0.2f (+/- %0.3f)" % (nb_binary_scores.mean(), nb_binary_scores.std() * 2))
```

<!-- grade_id: cell-6a24d3fda9d1c9d3, points: 5, grade: true, solution: false -->
```python
## Do not edit this cell - autograder test
assert(round(lr_binary_scores.mean(),3) == 0.865)
assert(round(nb_binary_scores.mean(), 3) == 0.83)
print("visable test cases passed!")
```

**(d)** (25 points) Let's also add some features. Specifically, we will add the count of part-of-speech (POS): for a particular document, how many nouns are there? How many verbs?... and so on.

* **(d-i)** write a function `get_pos_tags(text)` that when given a document, returns POS tag counter. Use NLTK for this.

<!-- grade_id: cell-bebd5ef221f28c34, points: 0, grade: false, solution: true -->
```python
%%time
import nltk
from collections import Counter

def get_pos_tags(text) -> Counter:
    """ when given a string, returns a POS tag counter, using NLTK"""
    text = str(text)
    # YOUR CODE HERE
    sent_tokens = nltk.sent_tokenize(text)
    sent_tokens = [nltk.word_tokenize(sentence) for sentence in sent_tokens]
    pos_sentences = nltk.pos_tag_sents(sent_tokens);
    
    
    concatenated = []
    for sentence in pos_sentences:
        concatenated = concatenated + sentence
    
    return Counter([item[1] for item in concatenated])


counter = get_pos_tags("This is a test case with a two singular nouns")
```

<!-- grade_id: cell-dffa39a8d4b55b08, points: 5, grade: true, solution: false -->
```python
%%time
## Do not edit this cell - autograder test
assert(get_pos_tags("This is a test case with a two singular nouns.")["NN"] == 2)
assert(get_pos_tags("This is a test case with multiple sentences. Will you get it right?")["MD"] == 1)
assert(get_pos_tags(str(docs_train[0]))["JJ"] == 98)
print("visable test cases passed!")
```

* **(d-ii)** now, write a function `get_pos_features` to create a dataframe for the POS features.

<!-- grade_id: cell-0da2866a5e6d384b, points: 0, grade: false, solution: true -->
```python
%%time
import pandas as pd

def get_pos_features(docs: list) -> pd.DataFrame:
    
    pos_docs = [get_pos_tags(x) for x in docs]
    rows = [x for x in range(len(docs))]
    panda = pd.DataFrame(pos_docs,rows)
    
    return panda.fillna(0)

print(get_pos_features(docs_train[:2])["NNP"])
get_pos_features(docs_train[:2]).head(3)
```

<!-- grade_id: cell-350dd0be2e32678d, points: 5, grade: true, solution: false -->
```python
%%time
## Do not edit this cell - autograder test
assert(get_pos_features(docs_train[:2])["JJ"][0] == 98)
# make sure you take care of any nans and make them 0! 
# this will be necessary for later!
assert(get_pos_features(docs_train[:2])["NNP"][1] == 0) 
print("visable test cases passed!")
```

```python
%%time
import pickle 

# # run get_pos_features on all the data
# pos_features = get_pos_features(dataset.data)

# # pickle pos_features
# with open('pos_features.pickle', 'wb') as f:
#     pickle.dump(pos_features, f)

#load pickled pos_features
with open('pos_features.pickle', 'rb') as f:
     pos_features = pickle.load(f)

```

* **(d-ii)** now, write a function `combine_data` and use ColumnTransformer to combine *unigram* and POS features. Note, using ColumnTransformer will return a numpy array with a smaller amount of data, but a scipi sparse matrix with more data.  

<!-- grade_id: cell-51b60e6c8888218b, points: 0, grade: false, solution: true -->
```python
%%time
# import pandas as pd
from sklearn.compose import ColumnTransformer
from sklearn.preprocessing import OneHotEncoder
import scipy

def transform_data(raw_data) -> pd.DataFrame:

    # if it's the full dataset, load the pickle!
    if len(raw_data) == 1600:
        print("loading pos from pickle")
        with open('pos_features.pickle', 'rb') as f:
             df_pos = pickle.load(f)
                
    # otherwise get the features for the data
    else:  
        df_pos = get_pos_features(raw_data)
        
    ## combine into a dataframe
    df_pos.insert(1,"docs", raw_data)
    ##### column transformer
    column_trans = ColumnTransformer(
    [("words", CountVectorizer(), "docs")],
        remainder='passthrough'
    )
    
    X_train = column_trans.fit_transform(df_pos)
        
    return X_train




#stuff = get_unigram_features(docs_train[:3]).toarray()
print(transform_data(docs_train[:3]))

```

<!-- grade_id: cell-3fc92c39b371c522, points: 10, grade: true, solution: false -->
```python
## Do not edit this cell - autograder test
assert(type(transform_data(docs_train[:3])) == np.ndarray)
assert(transform_data(docs_train[:3])[0][5] == 4)
assert(transform_data(docs_train[:3])[1][740] == 3)
assert(type(transform_data(docs_train[:20])) == scipy.sparse.csr.csr_matrix)
print("visable test cases passed!")
```

```python
%%time
# Since it took me 14 minutes to run transform_data  on all the data, 
# I have pickled my version for you! You will still need to be able to pass the tests on the first three 
import pickle 

# # run transform_data on entire dataset
# transformed_data = transform_data(dataset.data)

# # pickle transformed_data
# with open('transformed_data.pickle', 'wb') as f:
#     pickle.dump(transformed_data, f)


# load pickle
with open('transformed_data.pickle', 'rb') as f:
     transformed_data = pickle.load(f)

```

* **(d-iii)** finally, re-train your logistic regression model over the combined features. What is the 10-fold cross validation accuracy? (Make sure to use the same folds as before, so the results are comparable!)

<!-- grade_id: cell-04c93ab3bea940b9, points: 0, grade: false, solution: true -->
```python
%%time
#use random_state = 3


logmodel = LogisticRegression()
lr_xval_pos_scores = cross_val_score(logmodel, transformed_data, yy_train, cv=10)

print("Accuracy for Naive Bayes: %0.3f (+/- %0.3f)" % (lr_xval_pos_scores.mean(), lr_xval_pos_scores.std() * 2))
```

<!-- grade_id: cell-c34967bf011f9541, points: 5, grade: true, solution: false -->
```python
# do not edit this cell -- autograder
assert(round(lr_xval_pos_scores.mean(), 3) == 0.84) 
print("visable test cases passed!")
```

* **(d-iv)** Finally, some analysis: essentially, we have trained two models: LR with unigrams and bigrams, and LR with unigrams and POS tags. This setting allows us to compare the power of the new features (bigrams/POS tags). Which one is the better feature for the sentiment analysis task? Explain why this may be true. 

<!-- grade_id: cell-fd9caefd6217c3b2, points: 5, grade: true, solution: true -->
It seems that the better model is the one that containes the unigrams and the bigrams. This is because it did better during corss validation than the other model which means that it would probably better generalize the data. 